# Practice Midterm 1

## Question 1
Compute 
$$
    \sum_{i=1}^{n} \frac{1}{i} \sin\frac{i\pi}{n},
$$
when $n=1000$.
Store the result in the variable `partialSum`

<!-- grade_id: cell-d9d9c4d2c749a0e2, points: 0, grade: false, solution: true -->
```python
import math
# YOUR CODE HERE
raise NotImplementedError()
print(partialSum)
```

<!-- grade_id: cell-a42c9341abd4a64f, points: 15, grade: true, solution: false -->
```python
import numpy as np
assert np.isclose(partialSum, 1.8503659938562658)
```

## Question 2

You are given a list `L` of numbers. Compute the minimum value `Lmin` and maximum values `Lmax` of all elements in `L`. Do not use python built-in min or max functions.

The cell below will compute a new list every time it is run. You do not need to understand its content or alter it in any way.

```python
import numpy as np
n = np.random.randint(low = 0, high = 21)
L = np.random.random(n)
print(L)
```

<!-- grade_id: cell-38a8da6eda1113f3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dc5be6a7c8ef2cc8, points: 20, grade: true, solution: false -->
```python
import numpy as np
assert np.isclose(Lmin,min(L))
assert np.isclose(Lmax,max(L))
```

## Question 3
The Cobb-Douglass equation in economics states that 
$$
Y(L,K) = A L^\alpha K^\beta
$$
where $A$, $L$, $K$, and $Y$ are some quantities and $\alpha$ and $\beta$ are some parameters. It does not really matter what these are or what they represent...
You are given values of $A$, $L$, $K$, $Y$, $\alpha$, and $\beta$. Write a python expression that computes $Y$.
As before the values of the parameters will change every time you run the cell below

```python
import numpy as np
A = np.random.random()
L = np.random.random()
K = np.random.random()
alpha = np.random.random()
beta = np.random.random()
```

<!-- grade_id: cell-594e7aeab0175333, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dca45087069a769e, points: 15, grade: true, solution: false -->
```python
# Your code is being tested, but you are not shown the tests
assert not Y == None
```

## Question 4
A geometric sequence with first term $a$ and ratio $r$ is defined as:
$$
a_0 = a,
$$
and 
$$
a_{n+1} = r.a_n.
$$

You are given random values of $a$, $r$, and $n$. Write a python code that computes $a_n$ and store the result in the variable `an` using loops. Do not use the direct formula for $a_n$ as a fuction of $a$, $r$, and $n$ even if you know it...

```python
import numpy as np
a = np.random.uniform(low = -10, high = 10)
r = np.random.uniform(low = 0.75, high = 1.25)
n = np.random.randint(low = 10, high = 101)
print(f'a: {a}')
print(f'r: {r}')
print(f'n: {n}')
```

<!-- grade_id: cell-5b1965fe3d0965cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print(an)
```

<!-- grade_id: cell-cbd6283ca0e73de7, points: 15, grade: true, solution: false -->
```python
# Your code is being tested, but you are not shown the tests
```

## Question 5
You are given a long list of random integer numbers `L`. Find the index `idx` of the first number in `L` that is a multiple of 7 but not of 3.

For instance, if `L=[1, 35, 21]`, `idx` should be 1.

```python
L = np.random.randint(low = 0, high = 1000, size = 100)
L[-1] = 14
```

<!-- grade_id: cell-2d6e06bb2ab94473, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-93f3f1db0c336495, points: 20, grade: true, solution: false -->
```python
assert L[idx]%7 == 0 and L[idx]%3 != 0
for i in range(idx):
    assert L[i]%3 == 0 or L[i]%7 != 0
```

## Question 6

You are given a string `paragraph` consisting of a full paragraph taken from Mary Shelley's "Frankenstein" and a  letter `l` of the alphabet. Compute the number of times `n` at which the letter `l` appears in your paragraph. Store the result in the variable `nl`

You are not allowed to use the function `count`.


```python
# This commands extracts one paragraph from a file containing the entire novel
with open("Frankenstein.txt", "r") as f:
    Franky = [par.replace("\n", " ").lower() for par in f.read().split("\n\n") if len(par) > 50 and not par.startswith(" Letter")]
paragraph = Franky[np.random.randint(len(Franky))]
print(f"paragraph: \n   {paragraph}")
l = "qazwsxedcvfrtgbnhyujmkilop"[np.random.randint(26)]
print(f"letter: {l}")
```

<!-- grade_id: cell-423ace84e1fa32dc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a50d2d3723ddcdc6, points: 15, grade: true, solution: false -->
```python
assert nl == paragraph.count(l)
```