## Instructions
* Fill in your name, macID, and student number between quotes below.
* Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".
* The use of generative AI (ChatGPT, Claude, CoPilot, etc) is not allowed for this homework. The instructors will ask students suspected of using generative AI to go through the code line by line and explain it.
* Code is a form of communication and is highly personal. It is easy to spot similar assignments.
* You may remove any lines of the form `raise NotImplementedError()`.

```python
NAME =  ""
MACID = ""
STUDENT_NUMBER = ""
```

<!-- grade_id: cell-fad5c5f8ded68601, points: 0, grade: true, solution: false -->
```python
assert NAME != ""
assert MACID != ""
assert STUDENT_NUMBER != ""
```

# Assignment 1

## Question 1

Create a variable `i` and assign it the value 2026

<!-- grade_id: cell-2716e35f29836322, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4848caf7752b57be, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
assert type(i) == int
assert i == 2026
```

## Question 2
Increment `i` by 1

<!-- grade_id: cell-c077c89495fefafa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3d844ed7dde4c7dc, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
assert i == 2027
```

## Question 3
Compute the remainder of the division of `i` by 10

<!-- grade_id: cell-8391bc1ce7fd6c35, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Question 4

Divide `i` by 3 and store the result in a variable named `j` without altering the value of the variable `i`.

```python
#YOUR CODE HERE
raise NotImplementedError()
```

# Question 5
Round `i` to the **nearest** integer.

```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 6

Explain why the print statements in the following code print different texts

```python
theAnswer = "42"
print(theAnswer)
print("theAnswer")
```

<!-- grade_id: cell-42f70284d1824666, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 7

Explain why the following code generates an error.

```
alpha = 1
beta  = 2.3
mu    = alpha + beta
lambda = 1 - mu
sigma = alpha
```

<!-- grade_id: cell-4e950f50fd8bcea2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE