## Instructions
* Fill in your name, macID, and student number between quotes below.
* Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Delete any line containing the text `raise NotImplementedError()`
* Before submitting your assignments, make sure that your code runs without errors:
   First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All). Notebooks that generate errors will not receive and marks.
* As stated in the course outline, the use of generative AI (ChatGPT, Claude, CoPilot, etc) is not allowed for this homework. The instructors will ask students suspected of using generative AI to go through the code line by line and explain it.
* Code is a form of communication and is highly personal. It is easy to spot similar assignments.

```python
NAME =  ""
MACID = ""
STUDENT_NUMBER = ""
```

# Assignment 3

## Question 1 (2 marks): Geometric sequences


The geometric series with first term $a$ and ratio $n$ is defined as:
$$a_0 = a,\ a_{n+1} = r a_n.$$
Write a **recursive** function `geometricSequence` computing $a_n$ given $a$, $r$ and $n$.

Do not use the formula for the generic term of a geometric sequence.

<!-- grade_id: cell-85d110e9952b7e28, points: 0, grade: false, solution: true -->
```python
def geometricSequence(a: float, r: float, n: int) -> float:
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4848caf7752b57be, points: 2, grade: true, solution: false -->
```python
assert geometricSequence(2,0.5,1) == 1.0
assert geometricSequence(2,0.5,16) == 3.0517578125e-05
assert geometricSequence(2.2,1.6,4) == 14.417920000000004
```

## Question 2 (2 marks): The mid-point integration rule

The midpoint rule is special case of Riemann summ and can be used to approximate definte integrals.
It is given by
$$ 
    \int_a^b f(x)\, dx \simeq \sum_{i=1}^n \Delta{x} f(x_i^*),
$$
where $x_i^*$ is the mid-point of the interval $(x_{i-1}, x_i)$ where $x_i = a + \frac{i(b-a)}{n}$, and $\Delta{x} = \frac{b-a}{n}$.

Write a function `midPoint(f, a, b, n)` implementing the formula above.

<!-- grade_id: cell-9e9cddfa61660718, points: 0, grade: false, solution: true -->
```python
def midPoint(f, a, b, n):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0b0fe1ebf4e69a54, points: 2, grade: true, solution: false -->
```python
import numpy as np

def test1(x):
    return 1

def test2(x):
    return x

def test3(x):
    return np.exp(x)

# Test 1: integrating the constant function 1 on a interval (a,b) should return (b-a) up to machine precision
from hashlib import sha1
assert str(type(midPoint(test1, 1, 2, 11))) == """<class 'float'>""", """type of midPoint(test1, 1, 2, 11) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"""
assert str(round(midPoint(test1, 1, 2, 11), 2)) == """1.0""", """value of midPoint(test1, 1, 2, 11) is not correct (rounded to 2 decimal places)"""


# Test 2: when f(x) = x, a = 0, b = 1, the midpoint rule should approach 1/2 as n becomes large
assert str(type(midPoint(test2, 0, 1, 101))) == """<class 'float'>""", """type of midPoint(test2, 0, 1, 101) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"""
assert str(round(midPoint(test2, 0, 1, 101), 2)) == """0.5""", """value of midPoint(test2, 0, 1, 101) is not correct (rounded to 2 decimal places)"""

assert str(type(midPoint(test2, 0, 1, 10001))) == """<class 'float'>""", """type of midPoint(test2, 0, 1, 10001) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"""
assert str(round(midPoint(test2, 0, 1, 10001), 2)) == """0.5""", """value of midPoint(test2, 0, 1, 10001) is not correct (rounded to 2 decimal places)"""

assert str(type(midPoint(test2, 0, 1, 1000001))) == """<class 'float'>""", """type of midPoint(test2, 0, 1, 1000001) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"""
assert str(round(midPoint(test2, 0, 1, 1000001), 2)) == """0.5""", """value of midPoint(test2, 0, 1, 1000001) is not correct (rounded to 2 decimal places)"""


# Test 3: integrating e^x on (-1, 1) should give a result close to e - 1/e
assert str(type(midPoint(test3, -1, 1, 101))) == """<class 'numpy.float64'>""", """type of midPoint(test3, -1, 1, 101) is not correct"""
assert str(midPoint(test3, -1, 1, 101)) == """2.350363986223876""", """value of midPoint(test3, -1, 1, 101) is not correct"""


print('Success!')
```

# Question 3 (3 marks): Merge Sort

In lecture, we saw two sorting algorithms Bubble sort and Quicksort. Another popular sorting algorithm is called **Merge sort**. The algorithm can be described as follows:
- Start with an $n$-element unsorted structure $l$.
- Break $l$ into $n$ sub-lists each containing one element (each one element sub-list is now sorted!)
- Repeatedly merge sub-lists into sorted sub-lists until only one sub-list remains. This is the sorted list.

For example, if we start with the list $[6,7,3,1,8,7,2,4]$, then in the first step we will break up the list into the following pieces:
- $[6]$, $[7]$, $[3]$, $[1]$, $[8]$, $[7]$, $[2]$, and $[4]$.
- Now we start merging to make sorted sub-lists
- $[6,7]$, $[1,3]$, $[7,8]$, $[2,4]$
- $[1,3,6,7]$, $[2,4,7,8]$
- $[1,2,3,4,6,7,7,8]$

Write a function `merge(l)` which performs the above algorithm on an input list `l` and returns a sorted list.

```python
def merge(l : list) -> list:
    # YOUR CODE HERE
    return None
```

```python
assert merge([6,7,3,1,8,7,2,4]) == [1,2,3,4,6,7,7,8]
assert merge([]) == []
assert merge([1]) == [1]
assert merge([1,2,3,4,5,6,7,8]) == [1,2,3,4,5,6,7,8]
assert merge(list(range(0,100)).reverse) == list(range(0,100))
```

## Question 4 (2 marks): Pascal's triangle
Given two real numbers $a$ and $b$, and a non-negative integer $n$, one can extend the $n^{th}$ power of $a+b$ as:
$$
(a+b)^n = c_{n,0}a^n +c_{n,1}a^{n-1}b + \dots + c_{n,p}a^{n-1}b^b + \dots + c_{n,n}b^n,
$$
where the coefficients $c_{n,p}$ are given by the recursion formula
$$
  c_{n,p} = \begin{cases}
      1 & \text{ if } p = 0, \\
      c_{n-1,p-1} + c_{n-1,p} & \text{ if } 0 < p < n,\\
      1 & \text{ if } p = n.
  \end{cases}
$$
When writing the list $[c_{n,0}, c_{n,1}, \dots, c_{n,n}]$ one above each other, the result is a triangular list of numbers sometimes referred to as Pascal's triangle.

Write a function which returns a list containing the $n+1$ coefficients $[c_{n,0}, c_{n,1}, \dots, c_{n,n}]$ for any $n$.
You are not allowed to use the formula for $c_{n,p}$ if you know it. 

<!-- grade_id: cell-0ba9a356d3b0d4f4, points: 0, grade: false, solution: true -->
```python
def PascalTriangle(n):
    """ Compute the n^th row of  the Pascal triangle """
# YOUR CODE HERE
raise NotImplementedError()
for n in range(10):
    print(f"n: {n}, {PascalTriangle(n)}")
```

<!-- grade_id: cell-202fa94a027ab1e1, points: 2, grade: true, solution: false -->
```python
# Test 1 

# (a+b)^1 = a + b so c_{0,1} = c_{1,1} = 1
assert PascalTriangle(1) == [1, 1]
# (a+b)^3 = a^3 + 3 a^2 b + 3 a b^2 + b^3 so c_{0,3} = c_{3,3} = 1 and c{1,3} = c{2,3} = 3
assert PascalTriangle(3) == [1, 3, 3, 1]

# Test 2
from hashlib import sha1
assert str(type(PascalTriangle(5))) == """<class 'list'>""", """type of PascalTriangle(5) is not list. PascalTriangle(5) should be a list"""
assert str(len(PascalTriangle(5))) == """6""", """length of PascalTriangle(5) is not correct"""
assert str(sorted(map(str, PascalTriangle(5)))) == """['1', '1', '10', '10', '5', '5']""", """values of PascalTriangle(5) are not correct"""
assert str(PascalTriangle(5)) == """[1, 5, 10, 10, 5, 1]""", """order of elements of PascalTriangle(5) is not correct"""

assert str(type(PascalTriangle(7))) == """<class 'list'>""", """type of PascalTriangle(7) is not list. PascalTriangle(7) should be a list"""
assert str(len(PascalTriangle(7))) == """8""", """length of PascalTriangle(7) is not correct"""
assert str(sorted(map(str, PascalTriangle(7)))) == """['1', '1', '21', '21', '35', '35', '7', '7']""", """values of PascalTriangle(7) are not correct"""
assert str(PascalTriangle(7)) == """[1, 7, 21, 35, 35, 21, 7, 1]""", """order of elements of PascalTriangle(7) is not correct"""

assert str(type(PascalTriangle(9))) == """<class 'list'>""", """type of PascalTriangle(9) is not list. PascalTriangle(9) should be a list"""
assert str(len(PascalTriangle(9))) == """10""", """length of PascalTriangle(9) is not correct"""
assert str(sorted(map(str, PascalTriangle(9)))) == """['1', '1', '126', '126', '36', '36', '84', '84', '9', '9']""", """values of PascalTriangle(9) are not correct"""
assert str(PascalTriangle(9)) == """[1, 9, 36, 84, 126, 126, 84, 36, 9, 1]""", """order of elements of PascalTriangle(9) is not correct"""

assert str(type(PascalTriangle(15))) == """<class 'list'>""", """type of PascalTriangle(15) is not list. PascalTriangle(15) should be a list"""
assert str(len(PascalTriangle(15))) == """16""", """length of PascalTriangle(15) is not correct"""
assert str(sorted(map(str, PascalTriangle(15)))) == """['1', '1', '105', '105', '1365', '1365', '15', '15', '3003', '3003', '455', '455', '5005', '5005', '6435', '6435']""", """values of PascalTriangle(15) are not correct"""
assert str(PascalTriangle(15)) == """[1, 15, 105, 455, 1365, 3003, 5005, 6435, 6435, 5005, 3003, 1365, 455, 105, 15, 1]""", """order of elements of PascalTriangle(15) is not correct"""

print('Success!')
```

## Question 5 (2 marks) Polynomial evaluation

Consider a polynomial of degree $n$: $P(x) = a_0 + a_1x + a_2x^2 + \dots + a_n x^n$. $P$ is fully determined by its $n+1$ coefficients and can be represented by a python array of length $(n+1)$.

Write a function which given a list of coefficients `[a0, a1, ..., an]` and a float `x` returns $a_0 + a_1x + a_2x^2, +\dots + a_n x^n$. For instance, given `coefs = [1, 2, 3]` and `x = 2`, your function should return $1 + 2\cdot2 + 3\cdot 2^2 = 17$.

You are not allowed to use any implementation of polynomial evaluation from coefficients.

<!-- grade_id: cell-537481e373c9c73a, points: 0, grade: false, solution: true -->
```python
def polyEval(coefs: list[float, ...],  x:float) -> float:
    ''' Evaluate the polynomial defined by a list fo coefficients at a given point'''
# YOUR CODE HERE
raise NotImplementedError()
print(polyEval([1., 2., 3.], 2.))
```

<!-- grade_id: cell-c75869006b952fec, points: 2, grade: true, solution: false -->
```python
# Test 1 if a_0 = 1, P(0) = 1
assert polyEval([1, 2, 3, 4, 5], 0.) == 1.
assert polyEval([1, -22], 0.) == 1.

## Test 2: compare [1, 2, 3, -4] and 1 + 2x + 3x^2 - 4 x^3 at a few points
import numpy as np
P = [1, 2, 3, -4]
def PP(x):
    return 1.0 + 2. * x + 3. * x**2 - 4 * x**3
for x in (-2, 1, 2.23, 0., 4.):
    assert np.isclose(polyEval(P, x), PP(x))
```

## Question 6 (2 marks): Polynomial sum
Write a function which returns the sum of two polynomials `P` and `Q` represented by their coefficients. 
Your function should work even when `P` and `Q` don't have the same length. 

For instance, when `P = [1., 2., 3.]` and `Q = [-1., 2]`, your function should return the coefficient of the polynomial associated with $P(x) + Q(x) = 1 + 2x + 3x^2 -1 + 2x = 0 + 4x + 3$, that is `[0., 4., 3.]`

<!-- grade_id: cell-ef1f01225def90ba, points: 0, grade: false, solution: true -->
```python
def polySum(P: list[float, ...], Q: list[float, ...]) -> list[float, ...]:
    '''Returns the coefficients of P+Q'''
# YOUR CODE HERE
raise NotImplementedError()
print(polySum([1., 2., 3.], [-1., 2]))
```

<!-- grade_id: cell-30f188abb90be2ea, points: 2, grade: true, solution: false -->
```python
# Test 1:
assert polySum([1,], [-1, 2]) == [0, 2]
assert polySum([1, 2], [-1, 2]) == [0, 4]
assert polySum([1, 2], [1, 2, 3, 4, 5]) == [2, 4, 3, 4, 5]

# Test 2: generate 5 pairs of random polynomial and compare (P+Q)(x) and P(x) + Q(x) for
#         10 randomly generated values of x
import numpy as np
for i in range(5):
    p = np.random.randint(low = 1, high = 5)
    q = np.random.randint(low = 1, high = 5)
    P = list(np.random.uniform(low = -10, high = 10, size = p))
    Q = list(np.random.uniform(low = -10, high = 10, size = q))
    for j in range(10):
        x = np.random.uniform(low = -10, high = 10)
        assert np.isclose( polyEval(P,x) + polyEval(Q,x), polyEval(polySum(P, Q), x) ) 
```

## Question 7 (2 marks): Polynomial derivative

This one is easy...
Given `P`, compute the coefficients of its derivative. The only difficulty is that when computing the derivative of a constant, one should get `[0]` and not `[]`

<!-- grade_id: cell-fb5ed142a073e4f2, points: 0, grade: false, solution: true -->
```python
def polyDiff(P:list[float, ...]) -> list[float, ...]:
    ''' Compute the derivative of P'''
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-921f5c008b1f8ce5, points: 2, grade: true, solution: false -->
```python
assert polyDiff([3.]) == [0.]
assert polyDiff([1., 2., 1.]) == [2., 2.]
assert polyDiff([1., 1., 1., 1., 1.]) == [1., 2., 3., 4.]
```

## Question 8 (2 marks): Integer logarithm

Write a function `int_log(b,n)` which takes two integers $b$ and $n$ and computes the integer logarithm of $n$ base $b$. The integer logarithm of a number $n$ in base $b$ is the largest integer $m$, such that $b^m\leq n$. Your function should only use elementary Python operations, such as loops and basic arithmetic, you should use no external libraries or functions.

```python
def int_log(b,n): 
    # YOUR CODE HERE
    return None
```

```python
assert int_log(2,3) == 1
assert int_log(2,4) == 2
assert int_log(17,4914) == 3
```

## Question 9 (2 marks): $p$ norm
Another easy one to finish.
Write a function which given any $n$ numbers $a_0, a_1, \dots, a_{n-1}$ returns the $p$ norm:
$$
\|a_0, a_1, \dots, a_{n-1}\|_p = 
\left(\sum_{i=0}^{n-1} |a_i|^p\right)^{1/p}.
$$
Your function should be able to handle any number of arguments, and $p$ should be an optional argument with a default value of 2.

<!-- grade_id: cell-2dda9d0cb9536ac2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-127f865dea7f927c, points: 2, grade: true, solution: false -->
```python
# Test 1
assert pNorm(1.) == 1.
assert pNorm(1., 1., 1., 1.) == 2.
assert pNorm(4., p = 4) == 4.
assert pNorm(-4., p = 1) == 4.

from hashlib import sha1
assert str(type(pNorm(1., 2., 3.))) == """<class 'float'>""", """type of pNorm(1., 2., 3.) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"""
assert str(round(pNorm(1., 2., 3.), 2)) == """3.74""", """value of pNorm(1., 2., 3.) is not correct (rounded to 2 decimal places)"""

assert str(type(pNorm(1., 2., 3., p = 4))) == """<class 'float'>""", """type of pNorm(1., 2., 3., p = 4) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"""
assert str(round(pNorm(1., 2., 3., p = 4), 2)) == """3.15""", """value of pNorm(1., 2., 3., p = 4) is not correct (rounded to 2 decimal places)"""

assert str(type(pNorm(1., -2., 3., p = 4))) == """<class 'float'>""", """type of pNorm(1., -2., 3., p = 4) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"""
assert str(round(pNorm(1., -2., 3., p = 4), 2)) == """3.15""", """value of pNorm(1., -2., 3., p = 4) is not correct (rounded to 2 decimal places)"""

print('Success!')
```