## Instructions
* Fill in your name, macID, and student number between quotes below.
* If you @mcmaster.ca email address is `someone230@mcmaster.ca`, your macID is `aomeone230`. **Double check your nacID and student number or your assignment may not be graded.**
* Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Delete any line containing the text `raise NotImplementedError()`
* Before submitting your assignments, make sure that your code runs without errors:
   First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All). Notebooks that generate errors will not receive and marks.
* As stated in the course outline, the use of generative AI (ChatGPT, Claude, CoPilot, etc) is not allowed for this homework. The instructors will ask students suspected of using generative AI to go through the code line by line and explain it.
* Code is a form of communication and is highly personal. It is easy to spot similar assignments.

```python
NAME =  ""
MACID = ""
STUDENT_NUMBER = ""
```

# MATH 1MP3 Assignment 4

```python
import matplotlib.pyplot as plt
%matplotlib widget
import numpy as np
from ipywidgets import interact
```

## Question 1: The Intermediate Value Theorem (4 marks)

We recall the Intermediate Value Theorem from your first calculus class:

> Let $f$ be a continuous function over a closed interval $[a,b]$. For any $d$ betwwen $f(a)$ and $f(b)$, there is at least one $u$ in $[a,b]$ such that $f(c) = d$.

Use matplotlib to draw an *interactive* figure illustrating this theorem, similar to the one below (taken from the wikipedia entry for the theorem). 

<a title="Kpengboy, Public domain, via Wikimedia Commons" href="https://commons.wikimedia.org/wiki/File:Intermediatevaluetheorem.svg"><img width="500" alt="Intermediatevaluetheorem" src="https://upload.wikimedia.org/wikipedia/commons/thumb/e/e2/Intermediatevaluetheorem.svg/250px-Intermediatevaluetheorem.svg.png"></a>

You are free to use any function and any interval in your figure (provided that the IVT applies of course). Your figure should include a slider taking values between $f(a)$ and $f(b)$ to select $u$, the point $(c, f(c))$ should be shown in a contrasting colour and move as $u$ changes. 


<!-- grade_id: cell-3175901a677dc049, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 2: Taylor series (4 marks)

Another classical calculus concept is Taylor polynomials: Let $f$ be infinitely differentiable at a point $x_0$ and $n$ be a non-negative integer. The $n^\mathrm{th}$ order Taylor polynomial of $f$ at $x_0$ is:
$$
T_{n, x_0}(x) = \sum_{i=0}^n f^{(i)}(x_0)\frac{(x-x_0)^i}{i!},
$$
where $f^{(i)}$ is the $i^\mathrm{th}$ derivative of $f$.

Using the `numpy.polynomial` class, write a function `Taylor(x,n)` returning the Taylor polynomial or order $n$ of $f(x) = \sin x$ at some $x_0$. 

<!-- grade_id: cell-c6c582ae4b295133, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import numpy as np
import math

def deriv_sin(n):
    if n%4 == 0:
        return np.sin
    elif n%4 == 1:
        return np.cos
    elif n%4 == 2:
        return -np.sin
    else:
        return -np.cos


def Taylor(x,n, x0 = 0):
    coeff = []
    for i in range(n+1):
        coeff += [deriv_sin(i)(x0)*(x-x0)**2/math.factorial(i)]

```

<!-- grade_id: cell-21fef6bf43ac2347, points: 2, grade: true, solution: false -->
```python
### Test: Tn exist, takes a float and an in t and returns a polynomial
assert str(type(Tn)) == "<class 'function'>"
assert str(type(Tn(0., 5))) == """<class 'numpy.polynomial.polynomial.Polynomial'>"""
```

<!-- grade_id: cell-4d4d3258ceb33e9b, points: 2, grade: true, solution: false -->
```python
### Test: Tn is correct
from hashlib import sha1
assert str(type(Tn(0,3).coef)) == """<class 'numpy.ndarray'>""", """type of Tn(0,3).coef is not correct"""
assert str(Tn(0,3).coef) == """[ 0.          1.          0.         -0.16666667]""", """value of Tn(0,3).coef is not correct"""

assert str(type(Tn(np.pi/4, 3).coef)) == """<class 'numpy.ndarray'>""", """type of Tn(np.pi/4, 3).coef is not correct"""
assert str(Tn(np.pi/4, 3).coef) == """[-0.00924739  1.04437764 -0.07587321 -0.11785113]""", """value of Tn(np.pi/4, 3).coef is not correct"""

assert str(type(Tn(np.pi/3, 3).coef)) == """<class 'numpy.ndarray'>""", """type of Tn(np.pi/3, 3).coef is not correct"""
assert str(Tn(np.pi/3, 3).coef) == """[-0.03672655  1.132744   -0.17121331 -0.08333333]""", """value of Tn(np.pi/3, 3).coef is not correct"""

assert str(type(Tn(0,5).coef)) == """<class 'numpy.ndarray'>""", """type of Tn(0,5).coef is not correct"""
assert str(Tn(0,5).coef) == """[ 0.          1.          0.         -0.16666667  0.          0.00833333]""", """value of Tn(0,5).coef is not correct"""

assert str(type(Tn(np.pi/4, 5).coef)) == """<class 'numpy.ndarray'>""", """type of Tn(np.pi/4, 5).coef is not correct"""
assert str(Tn(np.pi/4, 5).coef) == """[ 2.02346334e-04  9.98492657e-01  4.62369655e-03 -1.74062940e-01\n  6.32276725e-03  5.89255651e-03]""", """value of Tn(np.pi/4, 5).coef is not correct"""

assert str(type(Tn(np.pi/3, 5).coef)) == """<class 'numpy.ndarray'>""", """type of Tn(np.pi/3, 5).coef is not correct"""
assert str(Tn(np.pi/3, 5).coef) == """[ 0.00142062  0.99204332  0.01836328 -0.18879067  0.01426778  0.00416667]""", """value of Tn(np.pi/3, 5).coef is not correct"""

print('Success!')
```

### Question 3 (4 marks)

Use the function you wrote for question 2 and create an interactive figure plotting $\sin x$ and $T_{n, x_0}(x)$. Your figure should use two sliders controlling respectively $x_0 \in (-\pi, pi)$ and $0 \le n \le 10$.


<!-- grade_id: cell-5672034530b4fcb9, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 4 (4 marks)

The *Legendre polynomials* are a family of polynomials defined by the recursive formula:

$$
\begin{align}
L_0(x) & = 1, \\
L_1(x) & = x, \\
L_{n+1}(x) & = \frac{2n+1}{n+1} x L_{n}(x) - \frac{n}{n+1} L_{n-1}(x).
\end{align}
$$

Write a function `Legendre(n)` returning the $n^\mathrm{th}$ Legendre polynomial using the numpy polynomial class. You are free to use recursivity or not. You are not allowed to use a function that computes Legendre polynomials directly.

<!-- grade_id: cell-68dbbcc186490565, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-401397ae836857d5, points: 2, grade: true, solution: false -->
```python
### Test: Legendre exists, takes 2 int and return an int
assert str(type(Legendre)) == "<class 'function'>"
assert str(type(Legendre(1))) == """<class 'numpy.polynomial.polynomial.Polynomial'>"""
```

<!-- grade_id: cell-6b106862a4797709, points: 2, grade: true, solution: false -->
```python
### Test: Legendre works

from hashlib import sha1
assert str(type(Legendre(5).coef)) == """<class 'numpy.ndarray'>""", """type of Legendre(5).coef is not correct"""
assert str(Legendre(5).coef) == """[ 0.     1.875  0.    -8.75   0.     7.875]""", """value of Legendre(5).coef is not correct"""

assert str(type(Legendre(10).coef)) == """<class 'numpy.ndarray'>""", """type of Legendre(10).coef is not correct"""
assert str(Legendre(10).coef) == """[-2.46093750e-01  0.00000000e+00  1.35351562e+01  0.00000000e+00\n -1.17304688e+02  0.00000000e+00  3.51914062e+02  0.00000000e+00\n -4.27324219e+02  0.00000000e+00  1.80425781e+02]""", """value of Legendre(10).coef is not correct"""

print('Success!')
```

### Question 5: Euclid's algorithm (4 marks)

Euclid's algorithm is a classical algorithm to compute the greatest common divisor (GCD) of two integers:

* let $a$ and $b$ be given
* set $a_0 = a$ and $b_0 = b$
* repeat:
    - $a_i = b_i  q_i + r_i$ with $r_i < b_i$.
    - if $r_0 = 0$ then return $GCD(a, b) = b_i$.
    - $a_{i+1} \leftarrow b_i$
    - $b_{i+1} \leftarrow r_i$.
* until $r_i = 0$

Note: since $b_{i+1} = r_i < b_i$, this algorithm will always converge in at most $b$ iterations (much less in practice)!  

Write a function `GCD` taking two `int` `a` and `b` and returning their greatesr common divider using the algorithm above. You are of course not alowed to use a python function that computes the GCD of two numbers, like `math.gcd`.

<!-- grade_id: cell-90e128663c21be0d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-7b5c98dbb27a8bb6, points: 2, grade: true, solution: false -->
```python
### Test: GCD exists, takes 2 int and return an int
assert str(type(GCD)) == "<class 'function'>"
assert str(type(GCD(1, 2))) == """<class 'int'>"""
```

<!-- grade_id: cell-db2209ab9d68cf8b, points: 2, grade: true, solution: false -->
```python
### Test: GCD works
import math
for i in range(10):
    a = np.random.randint(low = 1, high = 50)
    b = np.random.randint(low = 1, high = 50)
    assert GCD(a,b) == math.gcd(a,b)
```