## Instructions
* Fill in your name, macID, and student number between quotes below.
* If you @mcmaster.ca email address is `someone230@mcmaster.ca`, your macID is `aomeone230`. **Double check your nacID and student number or your assignment may not be graded.**
* Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Delete any line containing the text `raise NotImplementedError()`
* Before submitting your assignments, make sure that your code runs without errors:
   First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All). Notebooks that generate errors will not receive and marks.
* Submit your solution on avenue. **DO NOT RENAME THE NOTEBOOK FILES**. Notebooks that have been renamed will not be graded.
* As stated in the course outline, the use of generative AI (ChatGPT, Claude, CoPilot, etc) is not allowed for this homework. The instructors will ask students suspected of using generative AI to go through the code line by line and explain it.
* Code is a form of communication and is highly personal. It is easy to spot similar assignments.

```python
NAME =  ""
MACID = ""
STUDENT_NUMBER = ""
```

<!-- grade_id: cell-fad5c5f8ded68601, points: 0, grade: true, solution: false -->
```python
# The code below tries hard to check that the previous cell was filled corrrectly 
import re
assert NAME != ""
assert MACID != ""
assert STUDENT_NUMBER != ""
assert int(STUDENT_NUMBER)
assert re.search('[a-zA-Z]', MACID)
```

# Assignment 2

## Question 1

Consider the following text from the Python wikipedia page:

```python
text = "Python is a high-level, general-purpose programming language."
```

Using loops and conditional statements, create a new string named `txt` obtained from `text` by replacing each vowel in `text` by the character `*` if its index is **odd**. 

**Hint:** You can test if a letter is vowel by checking if it is in the string `"aeiouyAEIOUY"`

<!-- grade_id: cell-85d110e9952b7e28, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 2

Given an integer `n`, consider the following sequence: 
- If `n` is even, the next element in the sequence is `n` divided by two (this should still be an integer!)
- If `n` is odd, the next element in the sequence is `3n+1`.

Notice that once this sequence reaches `1` for the first time, it will oscillate endlessly: `1`, `4`, `2`, `1`,...

It is believed by many mathematicians that given any starting number this sequence always eventually reaches `1`. 

Create a list called `Collatz` consisting of the number of steps required for this sequence to first reach `1` for the first `100` integers (starting at `1`). 

For instance, your list should begin `[0,1,7,...]`

<!-- grade_id: cell-c077c89495fefafa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 3
Given the list `l` below, create two lists `lInt` and `lStr` consisting of respectively all integers in `l` and all strings in `l`. Use only one 'for' or 'while' loop.

```python
l = [1, 'two', 3, 4.0, 'cinq', 6, 7, 'EIGHT', 'IX', 10-1]
```

<!-- grade_id: cell-8391bc1ce7fd6c35, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-61715361ac67f6dc, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
assert lInt == [1, 3, 6, 7, 9]
assert lStr == ['two', 'cinq', 'EIGHT', 'IX']
```

## Question 4

The dot product of two vectors $a$ and $b$ of length $n$ is given by:
$$
a \cdot b = \sum_{i=0}^{n-1} a_i b_i.
$$
Use loops to compute the dot product of the vectors represented by the lists `a` and `b` below. Store the result in a variable named `dot`.

Note that every time you run the cell below, `a` abd `b` will change.

```python
# This cell generates random data. Do not worry if you do not understand how it works. 
# We will cover these functions later in the class
import numpy as np
n = 5
a = np.random.random(n)
b = np.random.random(n)

print(f'a: {a}')
print(f'b: {b}')
```

<!-- grade_id: cell-813da4ae93a56f7c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-41fb0a6cb4717919, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
assert np.isclose(dot, np.dot(a,b))
```

## Question 5
The *trace* of a $n \times n$ matrix $M$, noted $\mathrm{tr}\, M$ is the sum of its diagonal entries. For instance, the trace of the $3 \times 3$ matrix
$$M = \begin{bmatrix}
1 & -1 & 1\\
-2 & 2 & -1 \\
3 & 4 & 5
\end{bmatrix}$$
is $1 + 2 +5 = 8$.
Use *a single loop* to compute the trace of the $5 \times 5$ matrix $M$ below. Store the result in a variable named `trace`.

```python
# This cell generates random data. Do not worry if you do not understand how it works. 
# We will cover these functions later in the class
import numpy as np
n = 5
M = np.random.random(size = (5,5))
print(M)
```

<!-- grade_id: cell-9308725a3391f770, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3a042205d4ba2ea3, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
assert np.isclose(trace, np.trace(M))
```

# Question 6

The *transpose* of a matrix $M$, denoted by $M^T$, is defined to be the matrix whose rows are the columns of $M$. For example, 
$$\begin{bmatrix}
1 & -1 & 1\\
-2 & 2 & -1\\
3 & 4 & 5
\end{bmatrix}^T=\begin{bmatrix}
1 & -2 & 3\\
-1 & 2 & 4\\
1 & -1 & 5
\end{bmatrix}$$

A matrix $M$ is *symmetric* if $M^T=M$.

In the following cell you may assume that `m` is a variable of type `List` that has a nested list structure that represents some matrix (*i.e.* a nested list where each sublist has the same length). Write a nested list to create a Boolean variable `is_symmetric` that has value `True` if `m` represents a symmetric matrix and has value `False` otherwise.

```python
# a 3 by 3 matrix for testing purposes
m = [[1,2,3],[2,5,0],[3,0,7]]
```

<!-- grade_id: cell-c30b53cf0049b421, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5b599f0cc21ee6ff, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
assert is_symmetric == True
```

## Question 7
Consider the list `l` below. Use slicing to create a list 'lOdd' of all odd values in `l` sorted by increasing value (*i.e.* `lOdd = [1,3,...]`)

```python
l = [11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1, 0]
```

<!-- grade_id: cell-c87cc3131f1e281c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d058bba1a49ac0c0, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
assert lOdd == [1, 3, 5, 7, 9, 11]
```

## Question 8

Recall that an integer is said to be prime if its only factors are one and itself. An integer is composite if it is not prime. Suppose that the variable `n` has type `int` and is strictly positive. Write a loop which determines if `n` is prime and stores the answer to this question in a Boolean variable called `is_prime`. For example, if `n` has value `5`, then upon executing your loop, `is_prime` should have value `True`.

**Note:** one is not a prime number.

```python
# This cell generates random data. Do not worry if you do not understand how it works. 
# We will cover these functions later in the class
n = 2*int(50 * np.random.random())+1
print(f"n: {n}")
is_prime = True
```

<!-- grade_id: cell-5d1a2cee0505df21, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-13d2fe5242ec1195, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
import sympy as sp
assert is_prime == sp.isprime(n)
```

## Question 9

Recall that an integer is said to be prime if its only factors are one and itself. An integer is composite if it is not prime. Suppose that `l` a list consisting of positive integers. Using your work from **Question 8** create a new list called `primes` which contains all prime numbers contained in `l`. 

```python
l = [13,7,57,1013,2026,100001,502,2,1]
primes = []
```

<!-- grade_id: cell-640ba335863a02a9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a21500df45898acd, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
assert primes == [13,7,1013,2]
```

## Question 10
Given the strings `s1` and `s2` below, create a string `sMixed` consisting of the first character in `s1` followed by the first character in `s2`, the second character in `s1`, the second character in `s2` etc.    

```python
# This cell generates random data. Do not worry if you do not understand how it works. 
# We will cover these functions later in the class
import uuid
s1 = str(uuid.uuid4())
s2 = str(uuid.uuid4())
print(s1)
print(s2)
# s1 = 'This is a test'
# s2 = 'This is not it'
```

<!-- grade_id: cell-88e46b0cb229fa58, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4549cdfc8fee4eb4, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
assert sMixed[0::2] == s1
assert sMixed[1::2] == s2
```

## Question 11 

Recall that for vectors $v=(v_1,\dots,v_n)$ and $w=(w_1,\dots,w_n)$ of the same length, the dot product of $v$ and $w$, denoted by $v\cdot w$, can be computed as:
$$v\cdot w = v_1w_1+\cdots+v_nw_n $$

For example, if $v=(1,2,3)$ and $w=(4,5,6)$, then
$$v\cdot w = 1(4)+2(5)+3(6)= 4+10+18=32$$

For this question you will need the following two facts:
1. The norm (or length) of a vector $v$ can be computed as $|v|=\sqrt{v\cdot v}$.
2. If $\theta$ is the angle between vectors $v$ and $w$, then
$$ \mathrm{cos}(\theta)=\frac{v\cdot w}{|v||w|}$$

Given lists `t1` and `t2` of the same store the norms of `t1` and `t2` in variables named `norm1` and `norm2` and store the value of their dot product in a variable named `dot`, then compute the angle between `t1` and `t2` and store this value as a variable named `angle`. 

**Hint:** You may need some functions from Python's math library, but you may not use `numpy` `dot` function. 

```python
# This cell generates random data. Do not worry if you do not understand how it works. 
# We will cover these functions later in the class
t1 = np.random.random(size=3)
t2 = np.random.random(size=3)
print(f"t1: {t1}")
print(f"t2: {t2}")
```

<!-- grade_id: cell-2e7a2bc648458a92, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3e42f4f5be4efbdd, points: 2, grade: true, solution: false -->
```python
# Verification: this cell should run without errors.
import numpy as np
assert np.isclose(norm1, np.sqrt(np.dot(t1, t1)))
assert np.isclose(norm2, np.sqrt(np.dot(t2, t2)))
assert np.isclose(dot, np.dot(t1, t2))
assert np.isclose(dot, norm1 * norm2 * math.cos(angle))
```